# CMU auto-graded notebook

Before you turn these assignments in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE."

---

#CMU Machine Learning with Large Datasets

## Homework 1 - Part 1: Word Count

<!-- grade_id: collaborators, points: 0, grade: false, solution: true -->
```python
# Who did you collaborate with on this assignment? 
# if no one, collaborators should contain an empty string,
# else list your collaborators below

# collaborators = [""]
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_collaborators, points: 1, grade: true, solution: false -->
```python
try:
    collaborators
except:
    raise AssertionError("you did not list your collaborators, if any")
```

# ** 1. Word Count: Building a word count application**

This exercise will build on the techniques covered in the Spark tutorial to develop a simple word count application.  The volume of unstructured text in existence is growing dramatically, and Spark is an excellent tool for analyzing this type of data.  As a warm up, we will write code that calculates the most common words in the [Complete Works of William Shakespeare](http://www.gutenberg.org/ebooks/100) retrieved from [Project Gutenberg](http://www.gutenberg.org/wiki/Main_Page).

This could also be scaled to find the most common words in Wikipedia.

## During this word count question set we will cover:
* *Part 1:* Creating a base RDD and pair RDDs
* *Part 2:* Counting with pair RDDs
* *Part 3:* Finding unique words and a mean value
* *Part 4:* Apply word count to a file

> Note that for reference, you can look up the details of the relevant methods in:
> * [Spark's Python API](https://spark.apache.org/docs/latest/api/python/pyspark.html#pyspark.RDD)

## Part 1: Creating a base RDD and pair RDDs

In this part of the lab, we will explore creating a base RDD with `parallelize` and using pair RDDs to count words.

### (1a) Create a base RDD
We'll start by generating a base RDD by using a Python list and the `sc.parallelize` method.  Then we'll print out the type of the base RDD.

<!-- grade_id: sparkcontext_setup, points: 0, grade: true, solution: false -->
```python
# YOU CAN MOST LIKELY IGNORE THIS CELL. This is only of use for running this notebook locally.

# THIS CELL DOES NOT NEED TO BE RUN ON DATABRICKS. 
# Note that Databricks already creates a SparkContext for you, so this cell can be skipped.
import findspark
findspark.init()
import pyspark
from pyspark.sql import SQLContext
from nose.tools import assert_equal, assert_true
sc = pyspark.SparkContext(appName="hw")
sqlContext = SQLContext(sc)

print("spark context started")
```

```python
wordsList = ['cat', 'elephant', 'rat', 'rat', 'cat']
wordsRDD = sc.parallelize(wordsList, 4)
# Print out the type of wordsRDD
print(type(wordsRDD))
```

### (1b) Pluralize and test

Let's use a `map()` transformation to add the letter 's' to each string in the base RDD we just created. We'll define a Python function that returns the word with an 's' at the end of the word.  Please replace `<FILL IN>` with your solution. After you have defined `makePlural` you can run the cell which contains a test.  If you implementation is correct it will not print out anything; otherwise it will raise an error.

This is the general form that exercises will take, except that no example solution will be provided.  Exercises will include an explanation of what is expected, followed by code cells where one cell will have one or more `<FILL IN>` sections.  The cell that needs to be modified will have `# TODO: Replace <FILL IN> with appropriate code` on its first line.  Once the `<FILL IN>` sections are updated and the code is run, the test cell can then be run to verify the correctness of your solution.  The last code cell before the next markdown section will contain the tests.

<!-- grade_id: makePlural, points: 0, grade: false, solution: true -->
```python
# One way of completing the function
def makePlural(word):
    # TODO: Uncomment the template below and replace <FILL IN> with appropriate code
    # return <FILL IN>

    # YOUR CODE HERE
    raise NotImplementedError()

print(makePlural('cat'))
```

<!-- grade_id: test_makePlural, points: 1, grade: true, solution: false -->
```python
# Load in the testing code and check to see if your answer is correct
# If incorrect it will report back '1 test failed' for each failed test
# Make sure to rerun any cell you change before trying the test again

"""Check that makePlural function makes its input plural by adding an s"""
assert_equal(makePlural('rat'), 'rats')

```

### (1c) Apply `makePlural` to the base RDD

Now pass each item in the base RDD into a [map()](http://spark.apache.org/docs/latest/api/python/pyspark.html#pyspark.RDD.map) transformation that applies the `makePlural()` function to each element. And then call the [collect()](http://spark.apache.org/docs/latest/api/python/pyspark.html#pyspark.RDD.collect) action to see the transformed RDD.

<!-- grade_id: pluralRDD, points: 0, grade: false, solution: true -->
```python
# TODO: Uncomment the template below and replace <FILL IN> with appropriate code
# pluralRDD = wordsRDD.map(<FILL IN>)

# YOUR CODE HERE
raise NotImplementedError()

print(pluralRDD.collect())
```

<!-- grade_id: test_pluralRDD, points: 1, grade: true, solution: false -->
```python
"""Check that makePlural was applied to base RDD and call to collect returns correct output"""
assert_equal(pluralRDD.collect(), ['cats', 'elephants', 'rats', 'rats', 'cats'])
```

### (1d) Pass a `lambda` function to `map`

Let's create the same RDD using a `lambda` function.

<!-- grade_id: pluralLambdaRDD, points: 0, grade: false, solution: true -->
```python
# TODO: Uncomment the template below and replace <FILL IN> with appropriate code
# pluralLambdaRDD = wordsRDD.map(lambda <FILL IN>)

# YOUR CODE HERE
raise NotImplementedError()

print(pluralLambdaRDD.collect())
```

<!-- grade_id: test_pluralLambdaRDD, points: 1, grade: true, solution: false -->
```python
"""Check that lambda function applied to base RDD and call to collect returns correct output"""
assert_equal(pluralLambdaRDD.collect(), ['cats', 'elephants', 'rats', 'rats', 'cats'])

```

### (1e) Length of each word

Now use `map()` and a `lambda` function to return the number of characters in each word.  We'll `collect` this result directly into a variable.

<!-- grade_id: pluralLengths, points: 0, grade: false, solution: true -->
```python
# TODO: Uncomment the template below and replace <FILL IN> with appropriate code
# pluralLengths = (pluralRDD
#                  <FILL IN>
#                  .collect())

# YOUR CODE HERE
raise NotImplementedError()

print(pluralLengths)
```

<!-- grade_id: test_pluralLengths, points: 1, grade: true, solution: false -->
```python
"""Check that pluralLengths correctly computes the length of each word"""
assert_equal(pluralLengths, [4, 9, 4, 4, 4])
```

### (1f) Pair RDDs

The next step in writing our word counting program is to create a new type of RDD, called a pair RDD. A pair RDD is an RDD where each element is a pair tuple `(k, v)` where `k` is the key and `v` is the value. In this example, we will create a pair consisting of `('<word>', 1)` for each word element in the RDD.
We can create the pair RDD using the `map()` transformation with a `lambda()` function to create a new RDD.

<!-- grade_id: wordPairs, points: 0, grade: false, solution: true -->
```python
# TODO: Uncomment the template below and replace <FILL IN> with appropriate code
# wordPairs = wordsRDD.<FILL IN>

# YOUR CODE HERE
raise NotImplementedError()

print(wordPairs.collect())
```

<!-- grade_id: test_wordPairs, points: 1, grade: true, solution: false -->
```python
"""Check that wordPair contains pair RDDs containing (word, 1) pairs for a given input"""
assert_equal(wordPairs.collect(),
                  [('cat', 1), ('elephant', 1), ('rat', 1), ('rat', 1), ('cat', 1)])
```

## Part 2: Counting with pair RDDs

Now, let's count the number of times a particular word appears in the RDD. There are multiple ways to perform the counting, but some are much less efficient than others.

A naive approach would be to `collect()` all of the elements and count them in the driver program. While this approach could work for small datasets, we want an approach that will work for any size dataset including terabyte- or petabyte-sized datasets. In addition, performing all of the work in the driver program is slower than performing it in parallel in the workers. For these reasons, we will use data parallel operations.

### (2a) `groupByKey()` approach
An approach you might first consider (we'll see shortly that there are better ways) is based on using the [groupByKey()](http://spark.apache.org/docs/latest/api/python/pyspark.html#pyspark.RDD.groupByKey) transformation. As the name implies, the `groupByKey()` transformation groups all the elements of the RDD with the same key into a single list in one of the partitions.

There are two problems with using `groupByKey()`:
  + The operation requires a lot of data movement to move all the values into the appropriate partitions.
  + The lists can be very large. Consider a word count of English Wikipedia: the lists for common words (e.g., the, a, etc.) would be huge and could exhaust the available memory in a worker.

Use `groupByKey()` to generate a pair RDD of type `('word', iterator)`.

<!-- grade_id: wordsGrouped, points: 0, grade: false, solution: true -->
```python
# TODO: Uncomment the template below and replace <FILL IN> with appropriate code
# Note that groupByKey requires no parameters
# wordsGrouped = wordPairs.<FILL IN>

# YOUR CODE HERE
raise NotImplementedError()

for key, value in wordsGrouped.collect():
    print('{0}: {1}'.format(key, list(value)))
```

<!-- grade_id: test_wordsGrouped, points: 1, grade: true, solution: false -->
```python
"""Check that wordsGrouped contains words grouped with their counts for a given input"""
assert_equal(sorted(wordsGrouped.mapValues(lambda x: list(x)).collect()),
                  [('cat', [1, 1]), ('elephant', [1]), ('rat', [1, 1])])
```

### (2b) Obtaining the counts

Using the `groupByKey()` transformation creates an RDD containing 3 elements, each of which is a pair of a word and a Python iterator.

Now sum the iterator using a `map()` transformation.  The result should be a pair RDD consisting of (word, count) pairs.

<!-- grade_id: wordCountsGrouped, points: 0, grade: false, solution: true -->
```python
# TODO: Uncomment the template below and replace <FILL IN> with appropriate code
# wordCountsGrouped = wordsGrouped.<FILL IN>

# YOUR CODE HERE
raise NotImplementedError()

print(wordCountsGrouped.collect())
```

<!-- grade_id: test_wordCountsGrouped, points: 1, grade: true, solution: false -->
```python
"""Check that the sums of the groups are correct"""
assert_equal(sorted(wordCountsGrouped.collect()),
                  [('cat', 2), ('elephant', 1), ('rat', 2)])
```

### (2c) Counting using `reduceByKey` 

A better approach is to start from the pair RDD and then use the [reduceByKey()](http://spark.apache.org/docs/latest/api/python/pyspark.html#pyspark.RDD.reduceByKey) transformation to create a new pair RDD. The `reduceByKey()` transformation gathers together pairs that have the same key and applies the function provided to two values at a time, iteratively reducing all of the values to a single value. `reduceByKey()` operates by applying the function first within each partition on a per-key basis and then across the partitions, allowing it to scale efficiently to large datasets.

<!-- grade_id: wordCounts, points: 0, grade: false, solution: true -->
```python
# TODO: Uncomment the template below and replace <FILL IN> with appropriate code
# Note that reduceByKey takes in a function that accepts two values and returns a single value
# wordCounts = wordPairs.reduceByKey(<FILL IN>)

# YOUR CODE HERE
raise NotImplementedError()

print(wordCounts.collect())
```

<!-- grade_id: test_wordCounts, points: 1, grade: true, solution: false -->
```python
"""Check that use of reduceByKey for computing word count returns the correct values for a given input"""
assert_equal(sorted(wordCounts.collect()), [('cat', 2), ('elephant', 1), ('rat', 2)])
```

### (2d) All together

The expert version of the code performs the `map()` to pair RDD, `reduceByKey()` transformation, and `collect` in one statement.

<!-- grade_id: wordCountsCollected, points: 0, grade: false, solution: true -->
```python
# TODO: Uncomment the template below and replace <FILL IN> with appropriate code
# wordCountsCollected = (wordsRDD
#                        <FILL IN>
#                        .collect())

# YOUR CODE HERE
raise NotImplementedError()

print(wordCountsCollected)
```

<!-- grade_id: test_wordCountsCollected, points: 1, grade: true, solution: false -->
```python
"""Check that wordCountsCollected contains the correct (word, count) values"""
assert_equal(sorted(wordCountsCollected), [('cat', 2), ('elephant', 1), ('rat', 2)])

```

## Part 3: Finding unique words and a mean value

### (3a) Unique words

Calculate the number of unique words in `wordsRDD`.  You can use other RDDs that you have already created to make this easier.

<!-- grade_id: uniqueWords, points: 0, grade: false, solution: true -->
```python
# TODO: Uncomment the template below and replace <FILL IN> with appropriate code
# uniqueWords = <FILL IN>

# YOUR CODE HERE
raise NotImplementedError()

print(uniqueWords)
```

<!-- grade_id: test_uniqueWords, points: 0, grade: true, solution: false -->
```python
"""Check that uniqueWords equals the number of unique words in the given input"""
assert_equal(uniqueWords, 3)
```

### (3b) Mean using `reduce`

Find the mean number of words per unique word in `wordCounts`.

Use a `reduce()` action to sum the counts in `wordCounts` and then divide by the number of unique words.  First `map()` the pair RDD `wordCounts`, which consists of (key, value) pairs, to an RDD of values.

<!-- grade_id: mean, points: 0, grade: false, solution: true -->
```python
from operator import add

# TODO: Uncomment the template below and replace <FILL IN> with appropriate code
# totalCount = (wordCounts
#               .map(<FILL IN>)
#               .reduce(<FILL IN>))

# YOUR CODE HERE
raise NotImplementedError()

average = totalCount / float(wordCounts.count())
print(totalCount)
print(round(average, 2))
```

<!-- grade_id: test_mean, points: 2, grade: true, solution: false -->
```python
"""Check that totalCount contains the correct value"""
assert_equal(totalCount, 5)

"""Check that value contains the correct value"""
assert_equal(round(average, 2), 1.67)
```

## Part 4: Apply word count to a file

In this section we will finish developing our word count application.  We'll have to build the `wordCount` function, deal with real world problems like capitalization and punctuation, load in our data source, and compute the word count on the new data.

### (4a) `wordCount` function

First, define a function for word counting.  You should reuse the techniques that have been covered in earlier parts of this lab.  This function should take in an RDD that is a list of words like `wordsRDD` and return a pair RDD that has all of the words and their associated counts.

<!-- grade_id: def_wordCount, points: 0, grade: false, solution: true -->
```python
def wordCount(wordListRDD):
    """Creates a pair RDD with word counts from an RDD of words.

    Args:
        wordListRDD (RDD of str): An RDD consisting of words.

    Returns:
        RDD of (str, int): An RDD consisting of (word, count) tuples.
    """
    # TODO: Uncomment the template below and Replace <FILL IN> with appropriate code
    # return wordListRDD.<FILL IN>
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
print(wordCount(wordsRDD).collect())
```

<!-- grade_id: test_def_wordCount, points: 2, grade: true, solution: false -->
```python
"""Check that wordCount returns the correct output for a given input"""
assert_equal(sorted(wordCount(wordsRDD).collect()),
                  [('cat', 2), ('elephant', 1), ('rat', 2)])

```

### (4b) Capitalization and punctuation

Real world files are more complicated than the data we have been using in this lab. Some of the issues we have to address are:
  + Words should be counted independent of their capitialization (e.g., Spark and spark should be counted as the same word).
  + All punctuation should be removed.
  + Any leading or trailing spaces on a line should be removed.

Define the function `removePunctuation` that converts all text to lower case, removes any punctuation, and removes leading and trailing spaces.  Use the Python [re](https://docs.python.org/2/library/re.html) module to remove any text that is not a letter, number, or space.
If you are unfamiliar with regular expressions, you may want to review [this tutorial](https://developers.google.com/edu/python/regular-expressions) from Google.  Also, [this website](https://regex101.com/#python) is  a great resource for debugging your regular expression.

**Hints**

1. Use the [re.sub()](https://docs.python.org/2.7/library/re.html#re.sub) function.
2. For our purposes, "punctuation" means "not an alphabetic, numeric, or whitespace character." A convenient regular expression for matching a character that is not alpabetic, numeric, or whitespace is: `[^A-Za-z\s\d]`
3. Do _not_ use `\W`, as it retains underscores.

<!-- grade_id: removePunctuation, points: 0, grade: false, solution: true -->
```python
import re
def removePunctuation(text):
    """Removes punctuation, changes to lower case, and strips leading and trailing spaces.

    Note:
        Only spaces, letters, and numbers should be retained.  Other characters should should be
        eliminated (e.g. it's becomes its).  Leading and trailing spaces should be removed after
        punctuation is removed.

    Args:
        text (str): A string.

    Returns:
        str: The cleaned up string.
    """
    # TODO: Uncomment the template below and Replace <FILL IN> with appropriate code
    # return <FILL IN>
    
    # YOUR CODE HERE
    raise NotImplementedError()

print(removePunctuation('Hi, you!'))
print(removePunctuation(' No under_score!'))
print(removePunctuation(' *      Remove punctuation then spaces  * '))
```

<!-- grade_id: test_removePunctuation, points: 2, grade: true, solution: false -->
```python
"""Check that remotePunctuation removes capitalization and punctuation"""
assert_equal(removePunctuation(" The Elephant's 4 cats. "),
                  'the elephants 4 cats')

```

### (4c) Load a text file

For the next part of this lab, we will use the [Complete Works of William Shakespeare](http://www.gutenberg.org/ebooks/100) from [Project Gutenberg](http://www.gutenberg.org/wiki/Main_Page). To convert a text file into an RDD, we use the `SparkContext.textFile()` method. We also apply the recently defined `removePunctuation()` function using a `map()` transformation to strip out the punctuation and change all text to lower case.  Since the file is large we use `take(15)`, so that we only print 15 lines.

```python
# Just run this code

url = "https://raw.githubusercontent.com/10605/data/master/hw1/shakespeare.txt"

from pyspark import SparkFiles
sc.addFile(url)

shakespeareRDD = sc.textFile("file://" + SparkFiles.get("shakespeare.txt"), 8).map(removePunctuation)

print('\n'.join(shakespeareRDD
                .zipWithIndex()  # to (line, lineNum)
                .map(lambda kv: '{0}: {1}'.format(kv[1], kv[0]))  # to 'lineNum: line'
                .take(15)))
```

### (4d) Words from lines

Before we can use the `wordcount()` function, we have to address two issues with the format of the RDD:
  + The first issue is that  that we need to split each line by its spaces. ** Performed in (4d). **
  + The second issue is we need to filter out empty lines. ** Performed in (4e). **

Apply a transformation that will split each element of the RDD by its spaces. For each element of the RDD, you should apply Python's string [split()](https://docs.python.org/2/library/string.html#string.split) function. You might think that a `map()` transformation is the way to do this, but think about what the result of the `split()` function will be.

> Note:
> * Do not use the default implemenation of `split()`, but pass in a separator value.  For example, to split `line` by commas you would use `line.split(',')`.

<!-- grade_id: shakespeareWordsRDD, points: 0, grade: false, solution: true -->
```python
# hint: think about the difference between 'Map' and 'flatMap'
# TODO: Uncomment the template below and replace <FILL IN> with appropriate code
# shakespeareWordsRDD = shakespeareRDD.<FILL_IN>

# YOUR CODE HERE
raise NotImplementedError()

shakespeareWordCount = shakespeareWordsRDD.count()
print(shakespeareWordsRDD.top(5))
print(shakespeareWordCount)
```

<!-- grade_id: test_shakespeareWordsRDD, points: 2, grade: true, solution: false -->
```python
"""Check for the top 5 words (in descending order)"""
assert_equal(shakespeareWordsRDD.top(5), [u'zwaggerd', u'zounds', u'zounds', u'zounds', u'zounds'])

```

### (4e) Remove empty elements

The next step is to filter out the empty elements.  Remove all entries where the word is `''`.

<!-- grade_id: shakeWordsRDD, points: 0, grade: false, solution: true -->
```python
# TODO: Uncomment the template below and replace <FILL IN> with appropriate code
# shakeWordsRDD = shakespeareWordsRDD.<FILL_IN>

# YOUR CODE HERE
raise NotImplementedError()

shakeWordCount = shakeWordsRDD.count()
print(shakeWordCount)
```

<!-- grade_id: test_shakeWordsRDD, points: 2, grade: true, solution: false -->
```python
"""Check for the total number of words after removing empty elements"""
assert_equal(shakeWordCount, 882996)

```

### (4f) Count the words

We now have an RDD that is only words.  Next, let's apply the `wordCount()` function to produce a list of word counts. We can view the top 15 words by using the `takeOrdered()` action; however, since the elements of the RDD are pairs, we need a custom sort function that sorts using the value part of the pair.

You'll notice that many of the words are common English words. These are called stopwords. In a later lab, we will see how to eliminate them from the results.
Use the `wordCount()` function and `takeOrdered()` to obtain the fifteen most common words and their counts.

<!-- grade_id: top15WordsAndCounts, points: 0, grade: false, solution: true -->
```python
# TODO: Uncomment the template below and replace <FILL IN> with appropriate code
# top15WordsAndCounts = <FILL IN>

# YOUR CODE HERE
raise NotImplementedError()

print('\n'.join(map(lambda kv: '{0}: {1}'.format(kv[0], kv[1]), top15WordsAndCounts)))
```

<!-- grade_id: test_top15WordsAndCounts, points: 1, grade: true, solution: false -->
```python
"""Check for the top15 words and their counts"""
assert_equal(top15WordsAndCounts,
            [(u'the', 27361), (u'and', 26028), (u'i', 20681), (u'to', 19150), (u'of', 17463),
                   (u'a', 14593), (u'you', 13615), (u'my', 12481), (u'in', 10956), (u'that', 10890),
                   (u'is', 9134), (u'not', 8497), (u'with', 7771), (u'me', 7769), (u'it', 7678)])
```

## Export the notebook as an IPython notebook, then submit it to Gradescope!